<div class="alert alert-info">
    
➡️ Before you start, make sure that you are familiar with the **[study guide](https://liu-nlp.ai/text-mining/logistics/)**, in particular the rules around **cheating and plagiarism** (found in the course memo).

➡️ If you use code from external sources (e.g. StackOverflow, ChatGPT, ...) as part of your solutions, don't forget to add a reference to these source(s) (for example as a comment above your code).

➡️ Make sure you fill in all cells that say **`YOUR CODE HERE`** or **YOUR ANSWER HERE**.  You normally shouldn't need to modify any of the other cells.

</div>

# L1: Information Retrieval

In this lab you will apply basic techniques from information retrieval to implement the core of a minimalistic search engine. The data for this lab consists of a collection of app descriptions scraped from the [Google Play Store](https://play.google.com/store/apps?hl=en). From this collection, your search engine should retrieve those apps whose descriptions best match a given query under the vector space model.

```python
# Define some helper functions that are used in this notebook
from IPython.display import display, HTML

def success():
    display(HTML('<div class="alert alert-success"><strong>Checks have passed!</strong></div>'))
```

## Dataset

The app descriptions come in the form of a compressed [JSON](https://en.wikipedia.org/wiki/JSON) file. Start by loading this file into a [Pandas](https://pandas.pydata.org) [DataFrame](https://pandas.pydata.org/pandas-docs/stable/getting_started/dsintro.html#dataframe).

```python
import bz2
import numpy as np
import pandas as pd
pd.set_option('display.max_colwidth', 500)

with bz2.open('app-descriptions.json.bz2', mode='rt', encoding='utf-8') as source:
    df = pd.read_json(source, encoding='utf-8')
```

In Pandas, a DataFrame is a table with indexed rows and labelled columns of potentially different types. You can access data in a DataFrame in various ways, including by row and column. To give an example, the code in the next cell shows rows 200–204:

```python
df.loc[200:205]
```

Output:
```
name  \
200  Brick Breaker Star: Space King   
201      Brick Classic - Brick Game   
202     Bricks Breaker - Glow Balls   
203            Bricks Breaker Quest   
204             Brothers in Arms® 3   
205       Brown Dust - Tactical RPG   

                                                                                                                                                                                                                                                                                                                                                                                                                                                                                                             description  
200  Introducing the best Brick Breaker game that everyone can enjoy.\nEnjoy various missions and addictively simple play control.\n\n[Features]\n- Hundreds of stages and various missions\n- No limit to play such as Heart, play as much as you can!\n-
```

As you can see, there are two labelled columns: `name` (the name of the app) and `description` (a textual description). The next cell shows how to access only the description field from row 200:

```python
df.loc[200, 'description']
```

Output:
```
'Introducing the best Brick Breaker game that everyone can enjoy.\nEnjoy various missions and addictively simple play control.\n\n[Features]\n- Hundreds of stages and various missions\n- No limit to play such as Heart, play as much as you can!\n- 5 kinds of various items and items reinforcement system\n- No network required\n- game file is as low as 20M, light-weight download!\n- supports tablet screen\n- supports Google Play Leaderboards, Achievement, Multiplay\n- supports 14 languages\n\nHomepage:\nhttps://play.google.com/store/apps/dev?id=4931745640662708567\n\nFacebook: \nhttps://www.facebook.com/spcomesgames/'
```

## Problem 1: What's in a vector?

We start by vectorising the data — more specifically, we map each app description to a tf–idf vector. This is very simple with a library like [scikit-learn](https://scikit-learn.org/stable/), which provides a [TfidfVectorizer](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html) class for exactly this purpose.  If we instantiate this class, and call `fit_transform()` on all of our app descriptions, scikit-learn will preprocess and tokenize each app description, compute tf–idf values for each of them, and return a vectorised representation:

```python
from sklearn.feature_extraction.text import TfidfVectorizer

vectorizer = TfidfVectorizer()
X = vectorizer.fit_transform(df['description'])
X
```

Output:
```
<Compressed Sparse Row sparse matrix of dtype 'float64'
	with 267110 stored elements and shape (1614, 27877)>
```

Let’s pick the app "Pancake Tower", which has a rather short description text, to see how it has been vectorised:

```python
# We can use 'toarray' to convert the sparse matrix object into a "normal" array
vec = X[1032].toarray()[0]

# The app description & its corresponding vector
df.loc[1032, 'description'], vec
```

Output:
```
("Let's see how many pancakes you can pile up!!",
 array([0., 0., 0., ..., 0., 0., 0.], shape=(27877,)))
```

That's not very informative yet.  We know that the vector contains tf–idf values, and that each dimension of the vector corresponds to a token in the vectorizer’s vocabulary; let's extract these for this specific example.

Your **first task** is to find out how to access the `vectorizer`’s vocabulary, for example by [checking the documentation of `TfidfVectorizer`](https://scikit-learn.org/stable/modules/generated/sklearn.feature_extraction.text.TfidfVectorizer.html), and print all the tokens that are represented in the vector with a tf–idf value greater than zero (i.e., only the tokens that are actually part of this app’s description) _in descending order of the tf–idf values_.  In other words, the token with the highest tf–idf value should be at the top of your output, and the token with the lowest tf–idf value at the bottom.   Before you implement this, think about what you would expect the output look like, for example which words you would expect to have the highest/lowest tf–idf values in this example.

Your final output should look something like this:

```
<token 1>: <tf-idf value 1>
<token 2>: <tf-idf value 2>
...
```

<!-- grade_id: cell-95a3350b5c599c39, points: 0, grade: true, solution: true -->
```python
"""Print the tokens and their tf–idf values, in descending order."""

# YOUR CODE HERE
feature_names = vectorizer.get_feature_names_out()
print(f'feature names: {feature_names}')

dense = X[1032].toarray().flatten()
print(f'dense: {dense}')

mydf = pd.DataFrame({'term': feature_names, 'tfidf': dense})
mydf = mydf[mydf.tfidf > 0].sort_values('tfidf', ascending=False)
print(mydf)
```

Output:
```
feature names: ['00' '000' '0000' ... 'ﬁrst' 'ﬂip' 'ﬂying']
dense: [0. 0. 0. ... 0. 0. 0.]
           term     tfidf
15455  pancakes  0.653933
15954      pile  0.530470
12300       let  0.261529
18684       see  0.255763
13059      many  0.234920
10230       how  0.211532
22697        up  0.172168
3535        can  0.130476
24136       you  0.102769
```

## Problem 2: Finding the nearest vectors

To build a small search engine, we need to be able to turn _queries_ (for example the string "pile up pancakes") into _query vectors_, and then find out which of our app description vectors are closest to the query vector.

For the first part (turning queries into query vectors), we can simply re-use the `vectorizer` that we used for the app descriptions. For the second part, an easy way to find the closest vectors is to use scikit-learn’s [NearestNeighbors](https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.NearestNeighbors.html) class. This class needs to be _fit_ on a set of vectors (the "training set"; in our case the app descriptions) and can then be used with any vector to find its _nearest neighbors_ in the vector space.

**First,** instantiate and fit a class that returns the _ten (10)_ nearest neighbors:

<!-- grade_id: cell-bcebb6d4a62e7b7f, points: 0, grade: false, solution: true -->
```python
"""Instantiate and fit a class that returns the 10 nearest neighboring vectors."""

# YOUR CODE HERE
from sklearn.neighbors import NearestNeighbors
nn = NearestNeighbors(n_neighbors=10)
nn.fit(X=X)
```

Output:
```
NearestNeighbors(n_neighbors=10)
```

**Second,** implement a function that uses the vectorizer and the fitted class to find the nearest neighbours for a given query string:

<!-- grade_id: cell-18bf238a099d709b, points: 0, grade: false, solution: true -->
```python
def search(query):
    """Find the nearest neighbors in `df` for a query string.

    Arguments:
      query (str): A query string.

    Returns:
      The 10 apps (with name and description) most similar (in terms of
      cosine similarity) to the given query as a Pandas DataFrame.
    """
    
    # YOUR CODE HERE
  
    v_q = vectorizer.transform([query])
    ns = nn.kneighbors(v_q, return_distance=False)[0]
    return df.iloc[np.array(ns)]
    
  
# r = search('pile up pancakes')
# print(np.shape(r))
# print(r)
# print(np.(r)[:, 0, :])


```

### 🤞 Test your code

Test your implementation by running the following cell, which will sanity-check your return value and show the 10 best search results for the query _"pile up pancakes"_:

<!-- grade_id: cell-8506d7fe5961f87b, points: 1, grade: true, solution: false -->
```python
"""Check that searching for "pile up pancakes" returns a DataFrame with ten results,
   and that the top result is "Pancake Tower"."""

result = search('pile up pancakes')
display(result)
assert isinstance(result, pd.DataFrame), "search() function should return a Pandas DataFrame"
assert len(result) == 10, "search() function should return 10 search results"
assert result.iloc[0]["name"] == "Pancake Tower", "Top search result should be 'Pancake Tower'"
success()
```

Output:
```
name  \
1032                                      Pancake Tower   
326                     Cooking School: Games for Girls   
656   Hell’s Cooking — crazy chef burger, kitchen fever   
1235                                          Solitaire   
1164                                       Rummy - Free   
1181                       Sago Mini Trucks and Diggers   
436                         Dr. Panda's Ice Cream Truck   
1442                                    Turbo Dismount™   
1446                                              UNO!™   
1326                                     TO-FU Oh!SUSHI   

                                                                                                                                                                                                                                                                                                                                                                                                                 
```

Before continuing with the next problem, play around a bit with this simple search functionality by trying out different search queries, and see if the results look like what you would expect:

```python
# Example — try out your own queries!
search("christmas travel")
```

Output:
```
name  \
661                      Hey Duggee: The Tinsel Badge   
748                   Kids Animals Jigsaw Puzzles ❤️🦄   
310                     Coloring and drawing for kids   
848   MAGICA TRAVEL AGENCY – Free Match 3 Puzzle Game   
1208                                    Shuriken Jump   
311                            Coloring book for kids   
1357                             The Enchanted Worlds   
570                      Fruits Mania : Elly’s travel   
129                                 Basketball Battle   
633                     Gunspell - Match 3 Puzzle RPG   

                                                                                                                                                                                                                                                                                                                                                                                                                                     
```

## Problem 3: Custom preprocessing & tokenization

In Problem 1, you should have seen that `TfidfVectorizer` already performs some preprocessing by default and also does its own tokenization of the input data. This is great for getting started, but often we want to have more control over these steps. We can customize some aspects of the preprocessing through arguments when instantiating `TfidfVectorizer`, but for this exercise, we want to do _all_ of our preprocessing & tokenizing outside of scikit-learn.

Concretely, we want to use [spaCy](https://spacy.io), a library that we will make use of in later labs as well.  Here is a brief example of how to load and use a spaCy model:

```python
import spacy
# Load the small English model, disabling some components that we don't need right now
nlp = spacy.load('en_core_web_sm', disable=['parser', 'ner', 'textcat'])

# Take an example sentence and print every token from it separately
doc = nlp("Apple is looking at buying U.K. startup for $1 billion")
for token in doc:
    print(token.text)
```

Output:
```
Apple
is
looking
at
buying
U.K.
startup
for
$
1
billion
```

**Your task** is to write a preprocessing function that uses spaCy to perform the following steps:
- tokenization
- lemmatization
- stop word removal
- removing tokens containing non-alphabetical characters

We recommend that you go through the [Linguistic annotations](https://spacy.io/usage/spacy-101#annotations) section of the spaCy&nbsp;101, which demonstrates how you can get the relevant kind of information via the spaCy library.

Implement your preprocessor by completing the following function:

<!-- grade_id: cell-2df4eac94cdf6be3, points: 0, grade: false, solution: true -->
```python
def preprocess(text):
    """Preprocess the given text by tokenising it, removing any stop words, 
    replacing each remaining token with its lemma (base form), and discarding 
    all lemmas that contain non-alphabetical characters.

    Arguments:
      text (str): The text to preprocess.

    Returns:
      The list of remaining lemmas after preprocessing (represented as strings).
    """
    # YOUR CODE HERE
    doc = nlp(text)
    
    # print(f"{'TEXT':<15}{'LEMMA':<15}{'POS':<8}{'TAG':<10}{'DEP':<10}{'SHAPE':<10}{'IS_ALPHA':<10}{'IS_STOP':<10}{'CUSTOM_ALPHA':<10}")
    # print("-"*90)
    # for token in doc:
    #     print(f"{token.text:<15}{token.lemma_:<15}{token.pos_:<8}{token.tag_:<10}{token.dep_:<10}"
    #           f"{token.shape_:<10}{str(token.is_alpha):<10}{str(token.is_stop):<10}{str(is_alpha(token.text)):<10}")
    res = [token.lemma_ for token in doc if not token.is_stop and token.is_alpha]
    return res

    # raise NotImplementedError()
  
# doc = preprocess("Apple is looking at buying U.K. startup for $1 billion")
# print(doc)
```

Output:
```
['Apple', 'look', 'buy', 'startup', 'billion']
```

### 🤞 Test your code

Test your implementation by running the following cell:

<!-- grade_id: cell-642185b139f2cef6, points: 1, grade: true, solution: false -->
```python
"""Check that the preprocessing returns the correct output for a number of test cases."""

assert (
    preprocess('Apple is looking at buying U.K. startup for $1 billion') ==
    ['Apple', 'look', 'buy', 'startup', 'billion']
)
assert (
    preprocess('"Love Story" is a country pop song written and sung by Taylor Swift.') ==
    ['Love', 'Story', 'country', 'pop', 'song', 'write', 'sing', 'Taylor', 'Swift']
)
success()
```

Output:
```
<IPython.core.display.HTML object>
```

## Problem 4: The effect of preprocessing

To make use of the new `preprocess` function from Problem 3, we need to make sure that we incorporate it into `TfidfVectorizer` and disable all preprocessing & tokenization that `TfidfVectorizer` performs by default. Afterwards, we also need to re-fit the vectorizer and the nearest-neighbors class. To make this a bit easier to handle, let’s take everything we have done so far and put it in a single class `AppSearcher`.

### Task 4.1

**Your first task** is to complete the stub of the `AppSearcher` class given below. Keep in mind:
- The `fit()` function should fit both the vectorizer (from Problem 1) and the nearest-neighbors class (from Problem 2).  Make sure to modify the call to `TfidfVectorizer` to _disable all preprocessing & tokenization_ that it would do by default, and replace it with a call to the `preprocess()` function _defined in `AppSearcher`_.
- For the `preprocess()` function, you can start by copying your solution from Problem 3.
- For the `search()` function, you can copy your solution from Problem 2.
- Make sure to adapt your code to store the everything (data, vectorizer, nearest-neighbors class) within the `AppSearcher` class, so that your solution is independent of the code you wrote above!

<!-- grade_id: cell-cb53b61e9efe98af, points: 0, grade: true, solution: true -->
```python
class AppSearcher:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        self.vectorizer = TfidfVectorizer() 
        self.vectorizer = TfidfVectorizer(analyzer=self.preprocess)
        self.nn =  NearestNeighbors(n_neighbors=10)


        X = self.vectorizer.fit_transform(df['description'])
        self.nn.fit(X=X)
        # YOUR CODE HERE
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        doc = nlp(text)
       
        res = [token.lemma_ for token in doc if not token.is_stop and token.is_alpha]
        return res
        

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        v_q = self.vectorizer.transform([query])
        ns = self.nn.kneighbors(v_q, return_distance=False)[0]
        return self.df.iloc[np.array(ns)]

```

#### 🤞 Test your code

The following cell demonstrates how your class should be used. Note that it can take a bit longer to train it on the data as before, since we’re now calling spaCy for the preprocessing.

```python
apps = AppSearcher()
apps.fit(df)
apps.search("pile up pancakes")
```

Output:
```
name  \
1032                                      Pancake Tower   
326                     Cooking School: Games for Girls   
1235                                          Solitaire   
1181                       Sago Mini Trucks and Diggers   
1263                                   Spider Solitaire   
656   Hell’s Cooking — crazy chef burger, kitchen fever   
1164                                       Rummy - Free   
436                         Dr. Panda's Ice Cream Truck   
1245                                     Solitaire Free   
1442                                    Turbo Dismount™   

                                                                                                                                                                                                                                                                                                                                                                                                                 
```

### Task 4.2

**Your second task** is to experiment with the effect of using (or not using) different preprocessing steps.  We always need to _tokenize_ the text, but other preprocessing steps are optional and require a conscious decision whether to use them or not, such as:
- lemmatization
- lowercasing all characters
- removing stop words
- removing tokens containing non-alphabetical characters

**Modify the definition of the `preprocess()` function** of `AppSearcher` to include/exclude individual preprocessing steps, run some searches, and observe if and how the results change.  Which search queries you try out is up to you — you could compare searching for "pile up pancakes" with "pancake piling", for example; or you could try entirely different search queries aimed at different kinds of apps.  (You can modify the class directly by changing the cell above under Task 4.1, or copy the definitions to the cells below, whichever you prefer; there is no separate code to show for this task, but you will use your observations here for the individual reflection.)

```python
class AppSearcher2:
    def fit(self, df):
        """Instantiate and fit all the classes required for the search engine (cf. Problems 1 and 2)."""
        self.df = df
        self.vectorizer = TfidfVectorizer() 
        self.vectorizer = TfidfVectorizer(analyzer=self.preprocess)
        self.nn =  NearestNeighbors(n_neighbors=10)


        X = self.vectorizer.fit_transform(df['description'])
        self.nn.fit(X=X)
        # YOUR CODE HERE
        # raise NotImplementedError()

    def preprocess(self, text):
        """Preprocess the given text (cf. Problem 3)."""
        # YOUR CODE HERE
        doc = nlp(text)
       
        res = [token.text for token in doc if token.is_alpha and not token.is_stop]
        return res
        

    def search(self, query):
        """Find the nearest neighbors in `df` for a query string (cf. Problem 2)."""
        # YOUR CODE HERE
        v_q = self.vectorizer.transform([query])
        ns = self.nn.kneighbors(v_q, return_distance=False)[0]
        return self.df.iloc[np.array(ns)]

app2 = AppSearcher2()
app2.fit(df)


```

```python
print(apps.search("pancakes piling").iloc[0:10]['description'])

```

Output:
```
1032                                                                                                                                                                                                                                                                                                                                                                                                                                                                          Let's see how many pancakes you can pile up!!
326     Children like to help their parents. They especially like to help with cooking . When there is a cooking in the kitchen, it is no way to play. But cooking is a complicated process and often it ends up with a huge mess in the kitchen. But what if you are so eager to cook pancakes, cake or cupcakes? How to cook all that without doing a cleaning after? We have a solution! Home Cooking School with our curious Hippo has opened especially for parents and children! We do not only cook f
```

```python

print(app2.search("pancakes piling").iloc[0:10]['description'])
```

Output:
```
1032                                                                                                                                                                                                                                                                                                                                                                                                                                                                          Let's see how many pancakes you can pile up!!
326     Children like to help their parents. They especially like to help with cooking . When there is a cooking in the kitchen, it is no way to play. But cooking is a complicated process and often it ends up with a huge mess in the kitchen. But what if you are so eager to cook pancakes, cake or cupcakes? How to cook all that without doing a cleaning after? We have a solution! Home Cooking School with our curious Hippo has opened especially for parents and children! We do not only cook f
```

```python
print(apps.search("cooking").iloc[0:10]['description'])

```

Output:
```
825     Search for "BabyBus" for even more free panda games for you to try! \n\nEver wanted to have a taste of Chinese food? Now you can! BabyBus kids apps bring you a kitchen cooking game,that is, Panda Chef, Let's cook!, where you can cook within a multiple recipes. Pick ingredients, prepare food and explore cooking in your own way. You can also decorate your dishes with different toppings or condiments. Dress up as mini chefs, and clean up the messy kitchen! Don’t forget your chef’s hat!\n\nCooki...
326     Children like to help their parents. They especially like to help with cooking . When there is a cooking in the kitchen, it is no way to play. But cooking is a complicated process and often it ends up with a huge mess in the kitchen. But what if you are so eager to cook pancakes, cake or cupcakes? How to cook all that without doing a cleaning after? We have a solution! Home Cooking School with our curious Hippo has opened especially for parents and children! We do not only cook f
```

```python
print(app2.search("cooking").iloc[0:10]['description'])

```

Output:
```
324     Cooking Joy 2 - a new highly addictive cooking game from the team that brought you Cooking Joy, is calling all master chef candidates! Upgraded from Cooking Joy - a fun cooking game, it inherits the same challenging spirit and adds more fun! \n\nIf you have always dreamt of becoming a top chef in a crazy cooking game world, then catch the cooking crazy fever with this game! Download NOW and try it for FREE!\n\nTime to get back to the kitchen and enjoy cooking delicious dishes for starving cu...
323     Welcome to Cooking Hot👩‍🍳👨‍🍳, the BEST chef game, FAST PACE Cooking Game and Restaurant Game! It brings you the Purest and Truest COOKING experiences. Come to different cooking cities, feel the craze and madness of a cooking games lover. Cook and Serve food as a crazy chef! \n\nIn this chef game, you act as a top chef and star chef. The craze for the cooking grabs you, and you are cooking in brink of madness. You work in the kitchens of different countries and cooking styles, TAP
```

```python
print(apps.search("re-opening").iloc[0:10]['description'])

```

Output:
```
687     Imagine unlimited possibilities in your smart home. Set scenes and fast effects to your mood.\n\nExperience Dance Sensation in your entertainment area with Philips Hue Entertainment. Enjoy a more colorful ambiance on your IKEA TRADFRI gateway.\n\nFeel more in control with schedules and automation from sunrise to sunset. Widgets, Shortcuts, Quick settings tiles, and Wear OS help you get more out of your smart lights.\n\nControl multiple bridges simultaneously without switching between them.\n...
430     Dr. Panda Restaurant is re-opening, and this time all the choices are yours! Make the pizza of your dreams, a pasta dish to rave about, or a soup so spicy your customers will breathe fire! Sweet or salty? Spicy or bitter? It’s up to you!\n\nKids can take charge in their own kitchen in Dr. Panda Restaurant 2! Future fine chefs have the freedom to choose what they want to prepare and exactly how they’d like to prepare it! Chop, grate, blend, fry and more with over 20 ingredients an
```

```python
print(app2.search("re-opening").iloc[0:10]['description'])

```

Output:
```
687     Imagine unlimited possibilities in your smart home. Set scenes and fast effects to your mood.\n\nExperience Dance Sensation in your entertainment area with Philips Hue Entertainment. Enjoy a more colorful ambiance on your IKEA TRADFRI gateway.\n\nFeel more in control with schedules and automation from sunrise to sunset. Widgets, Shortcuts, Quick settings tiles, and Wear OS help you get more out of your smart lights.\n\nControl multiple bridges simultaneously without switching between them.\n...
957     Once upon a time, there was wannabe baker whose dream came true! Enjoy this mouth-watering cake bakery story and help Lizzie fulfill her dream of someday opening up a sweet bakery of her own. Now that day has finally arrived! She's graduating from college, and she's more than ready to get baking some tasty cupcakes. But she needs your help! Have fun opening up lots of yummy bakeries, baking with Lizzie and becoming a true baking professional who’s famous for beautiful, tasty deli
```

```python
str(df.loc[656]['description'])
```

Output:
```
"⭐ ⭐ ⭐ ⭐ ⭐ New world of crazy cooking is here. Feel what it means to be a master chef who prepares fantastic fast food in a prominent king kitchen! If you haven't ever tried yourself as a hamburger chef cook, it's possibly the best time for making diner. Download and launch Hell's Cooking — crazy chef burger, kitchen fever HD game and get prepared to jump into a fever and adventurous perfect world of burgers.\n\nNew girls game Hell's Cooking gives you lots of opportunities for your crazy cafe puzzle! Here you own a restaurant and rush serving numerous happy frenzy visitors from the street every day. Queen, you need to not only make perfect burger fast, tasty meals and bake delicious pancakes or cakes in chef role in cafe, but also serve hungry diner people in a dash. Stand and use everything at hand, as the whole kitchenette is at your full disposal. Here you will find whatever you want, as basic foodstuff for chef, so unique favors to make baking a fun! If your visitor orders a burger
```

## Individual reflection

<div class="alert alert-info">
    <strong>After you have solved the lab,</strong> write a <em>brief</em> reflection (max. one A4 page) on the question(s) below.  Remember:
    <ul>
        <li>You are encouraged to discuss this part with your lab partner, but you should each write up your reflection <strong>individually</strong>.</li>
        <li><strong>Do not put your answers in the notebook</strong>; upload them in the separate submission opportunity for the reflections on Lisam.</li>
    </ul>
</div>

1. In Problem 1, which token had the highest tf–idf score, which the lowest?  Based on your knowledge of how tf–idf works, how would you explain this result?
2. Based on your observations in Problem 4, which preprocessing steps do you think are the most appropriate for this "search engine" example?  Why?

**Congratulations on finishing this lab! 👍**

<div class="alert alert-info">
    
➡️ Before you submit, **make sure the notebook can be run from start to finish** without errors.  For this, _restart the kernel_ and _run all cells_ from top to bottom. In Jupyter Notebook version 7 or higher, you can do this via "Run$\rightarrow$Restart Kernel and Run All Cells..." in the menu (or the "⏩" button in the toolbar).

</div>

```python

```